[1501136562] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.002 second response time
[1501136622] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1501136682] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.012 second response time
[1501136713] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501136743] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1501136772] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501136832] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501136892] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501144790] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501144839] SERVICE ALERT: localhost;SSH;OK;SOFT;2;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1501144912] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501144959] SERVICE ALERT: localhost;SSH;OK;SOFT;2;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1501148390] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501148446] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 6.84, 3.10, 1.63
[1501148452] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1501148500] SERVICE ALERT: localhost;SSH;OK;SOFT;3;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1501148503] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 5.06, 3.31, 1.79
[1501148563] SERVICE ALERT: localhost;Current Load;OK;SOFT;3;OK - load average: 2.45, 2.89, 1.74
[1501148854] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501148903] SERVICE ALERT: localhost;SSH;OK;SOFT;2;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1501149046] SERVICE ALERT: ip-172-31-0-168;USERS;WARNING;SOFT;1;Critical [Warning]
[1501149104] SERVICE ALERT: ip-172-31-0-168;USERS;WARNING;SOFT;2;Critical [Warning]
[1501149164] SERVICE ALERT: ip-172-31-0-168;USERS;WARNING;SOFT;3;Critical [Warning]
[1501149229] SERVICE ALERT: ip-172-31-0-168;USERS;OK;SOFT;4;[OK]
[1501149335] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501149385] SERVICE ALERT: localhost;SSH;OK;SOFT;2;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1501149411] SERVICE ALERT: ip-172-31-0-168;USERS;WARNING;SOFT;1;Critical [Warning]
[1501149474] SERVICE ALERT: ip-172-31-0-168;USERS;OK;SOFT;2;[OK]
[1501150245] SERVICE ALERT: ip-172-31-0-168;USERS;CRITICAL;SOFT;1;[Critical]
[1501150304] SERVICE ALERT: ip-172-31-0-168;USERS;CRITICAL;SOFT;2;[Critical]
[1501150364] SERVICE ALERT: ip-172-31-0-168;USERS;CRITICAL;SOFT;3;[Critical]
[1501150424] SERVICE ALERT: ip-172-31-0-168;USERS;CRITICAL;HARD;4;[Critical]
[1501150664] SERVICE ALERT: ip-172-31-0-168;USERS;UNKNOWN;HARD;4;Remote command execution failed: sh: 1: /$/check_local_user_count.pl: not found
[1501150966] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501151023] SERVICE ALERT: localhost;SSH;OK;SOFT;2;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1501151145] SERVICE ALERT: ip-172-31-0-168;USERS;WARNING;HARD;4;[Warning]
[1501151455] SERVICE ALERT: ip-172-31-0-168;USERS;CRITICAL;HARD;4;CRITICAL - Plugin timed out after 10 seconds
[1501151456] SERVICE FLAPPING ALERT: ip-172-31-0-168;USERS;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1501151465] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501151509] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501151520] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 6.77, 3.47, 2.06
[1501151520] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 74 processes with STATE = RSZDT
[1501151569] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1501151586] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501151588] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 8.81, 4.67, 2.58
[1501151601] SERVICE ALERT: localhost;HTTP;CRITICAL;HARD;4;CRITICAL - Socket timeout after 10 seconds
[1501151644] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 11.54, 6.09, 3.18
[1501151663] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 74 processes with STATE = RSZDT
[1501151666] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;CRITICAL - Socket timeout after 10 seconds
[1501151691] SERVICE ALERT: ip-172-31-0-168;USERS;WARNING;HARD;4;[Warning]
[1501151706] SERVICE ALERT: localhost;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 10.43, 6.85, 3.63
[1501151727] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501151727] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;CRITICAL - Socket timeout after 10 seconds
[1501151727] SERVICE ALERT: localhost;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501151756] SERVICE ALERT: ip-172-31-0-168;USERS;CRITICAL;HARD;4;CRITICAL - Plugin timed out after 10 seconds
[1501151784] SERVICE ALERT: localhost;PING;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1501151787] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1501151838] SERVICE ALERT: localhost;Total Processes;OK;SOFT;3;PROCS OK: 74 processes with STATE = RSZDT
[1501151839] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 6.714 second response time
[1501151841] SERVICE ALERT: localhost;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.06 ms
[1501151902] SERVICE ALERT: localhost;HTTP;CRITICAL;HARD;4;CRITICAL - Socket timeout after 10 seconds
[1501152086] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501152136] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 4.577 second response time
[1501152136] SERVICE FLAPPING ALERT: localhost;HTTP;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1501152138] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 81 processes with STATE = RSZDT
[1501152142] SERVICE ALERT: localhost;SSH;OK;HARD;4;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1501152203] SERVICE ALERT: localhost;HTTP;CRITICAL;HARD;4;CRITICAL - Socket timeout after 10 seconds
[1501152206] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501152206] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1501152229] SERVICE ALERT: ip-172-31-0-168;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501152265] SERVICE ALERT: localhost;PING;UNKNOWN;SOFT;1;(No output on stdout) stderr:
[1501152268] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1501152268] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1501152286] SERVICE ALERT: ip-172-31-0-168;PING;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1501152324] SERVICE ALERT: localhost;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.06 ms
[1501152325] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;3;CRITICAL - Plugin timed out after 10 seconds
[1501152325] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;CRITICAL - Socket timeout after 10 seconds
[1501152346] SERVICE ALERT: ip-172-31-0-168;PING;CRITICAL;SOFT;3;CRITICAL - Plugin timed out after 10 seconds
[1501152378] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 7.248 second response time
[1501152382] SERVICE ALERT: localhost;Total Processes;OK;SOFT;4;PROCS OK: 76 processes with STATE = RSZDT
[1501152387] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;CRITICAL - Socket timeout after 10 seconds
[1501152408] SERVICE ALERT: ip-172-31-0-168;PING;CRITICAL;HARD;4;CRITICAL - Plugin timed out after 10 seconds
[1501152443] SERVICE ALERT: localhost;HTTP;CRITICAL;HARD;4;CRITICAL - Socket timeout after 10 seconds
[1501152447] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1501152505] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 81 processes with STATE = RSZDT
[1501152506] SERVICE ALERT: ip-172-31-0-168;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.49 ms